## Εργασία 2 ##

Καλωσήρθατε στην δεύτερη σας εργασία. Η εργασία αυτή έχει σκοπό να σας βοηθήσει να εμπεδώσετε τα δενδρικά μοντέλα και την αξιολόγηση μοντέλων.

Στην εργασία αυτή θα πρέπει να συμπληρώσετε κώδικα Python 3 στα σημεία που αναφέρουν # YOUR CODE HERE. Μην τροποποιείτε τον κώδικα που βρίσκεται εκτός αυτών των περιοχών.

Πρωτού παραδόσετε την εργασία σας σιγουρευτείτε ότι ο κώδικας σε όλα τα κελιά τρέχει σωστά. Για το σκοπό αυτό επιλέξτε από το μενού Χρόνος εκτέλεσης (runtime) -> Επανεκίνηση περιόδου λειτουργίας και εκτέλεση όλων.

Συμπληρώστε το όνομα (NAME) και το AEM σας παρακάτω:

```python
NAME = "Ιάσονας Κακανδρής"
AEM = "10003"
```

**1** Διαβάστε το διαθέσιμο από την sklearn σύνολο δεδομένων diabetes που αφορά την πρόβλεψη της πορείας που θα έχει ένας ασθενής με διαβήτη ένα χρόνο μετά (εργασία παλινδρόμησης) και αποθηκεύστε το σε μια μεταβλητή με το όνομα diabetes. (0.5 μονάδες)

<!-- grade_id: cell-ac5d304e50a110a6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.datasets import load_diabetes
diabetes = load_diabetes() 

```

<!-- grade_id: cell-2d158133c851eb4e, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθής ανάγνωσης του συνόλου δεδομένων"""
assert diabetes.feature_names == ['age', 'sex', 'bmi', 'bp', 's1', 's2', 's3', 's4', 's5', 's6']
```

**2** Χωρίστε το σύνολο που διαβάσατε σε υποσύνολο ελέγχου 20% και υποσύνολο εκπαίδευσης 80% με την μέθοδο train_test_split και σπόρο τυχαιότητας 42. Aποθηκεύστε το σύνολο εκπαίδευσης σε μεταβλητές X_train, y_train και το σύνολο ελέγχου σε μεταβλητές X_test, y_test. (0.5 μονάδες)

```python
#diabetes.target
#diabetes.data
diabetes.feature_names
```

Output:
```
['age', 'sex', 'bmi', 'bp', 's1', 's2', 's3', 's4', 's5', 's6']
```

<!-- grade_id: cell-a637aedb66ae6d89, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
"""
diabetes.data -> the independent parameters 
diabates.target -> the dependent parameter 
"""
from sklearn.model_selection import train_test_split
X_train , X_test , y_train , y_test = train_test_split( diabetes.data ,  diabetes.target , test_size= 0.2 , random_state= 42)
#raise NotImplementedError()
```

<!-- grade_id: cell-300331f01f617c3f, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == 0.02736
assert round(X_test[0][8], 5) == 0.03243
```

**3** Χρησιμοποιώντας τα δεδομένα εκπαίδευσης, εκπαιδεύστε δύο δένδρα παλινδρόμησης με κριτήρια διαχωρισμού τα squared_eror και poisson ([δείτε ενότητα 1.10.7.2](https://scikit-learn.org/stable/modules/tree.html#tree)) αντίστοιχα, θέτοντας τον σπόρο τυχαιότητας στην τιμή 42 και το μέγιστο βάθος στην τιμή 3. Υπολογίστε τη μετρική r2 με βάση τις προβλέψεις τους στα δεδομένα ελέγχου, τις οποίες αποθηκεύστε στις μεταβλητές y_pred1 (squared_error) και y_pred2 (poisson) αντίστοιχα.
Αποθηκεύστε τις μετρικές στις μεταβλητές r2_squared_error και r2_poisson αντίστοιχα. (1 μονάδα)

<!-- grade_id: cell-62cfb517c0f9271e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.tree import DecisionTreeRegressor 
from sklearn.metrics import r2_score
import numpy as np 

squared_error_tree = DecisionTreeRegressor(criterion='squared_error' , random_state=42 , max_depth=3) # for the squared error
poisson_tree = DecisionTreeRegressor(criterion='poisson' , random_state=42 , max_depth=3) # for the poisson 

# fit the data 
squared_error_tree.fit(X_train , y_train)
poisson_tree.fit(X_train , y_train) 

# predictions
y_pred1 = squared_error_tree.predict(X_test) 
y_pred2 = poisson_tree.predict(X_test)

# errors 
r2_squared_error = r2_score(y_true= y_test , y_pred= y_pred1)
r2_poisson  = r2_score(y_true= y_test , y_pred= y_pred2)

# print the values if we want
# print(r2_squared_error)
# print(r2_poisson)

```

<!-- grade_id: cell-714ce0b2b256307e, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού των μετρικών"""
assert round(r2_squared_error, 3) == 0.329
assert round(r2_poisson, 3) == 0.391
assert len(y_pred1) == len(y_pred2)
```

**4** Αναπτύξτε κώδικα με βάση την ιδιότητα [tree_](https://scikit-learn.org/stable/auto_examples/tree/plot_unveil_tree_structure.html#sphx-glr-auto-examples-tree-plot-unveil-tree-structure-py) του παρακάτω μοντέλου, προκειμένου να εμφανίσετε το μονοπάτι που ακολουθήθηκε στο δένδρο για την πρόβλεψη της πρώτης περίπτωσης του συνόλου ελέγχου. Σε περίπτωση που ελέγχεται δεύτερη φορά μια μεταβλητή στο μονοπάτι, θα πρέπει να διατηρούμε τον τελευταίο και στενότερο περιορισμό. Ακολουθούν παραδείγματα ορθής λειτουργίας, στα οποία φαίνεται και η μορφή αν-τότε που πρέπει να έχει το μονοπάτι. (4 μονάδες)

```python
from sklearn.tree import export_text


model = DecisionTreeRegressor(max_depth=3, random_state=42)
model.fit(X_train, y_train)


```

Output:
```
DecisionTreeRegressor(max_depth=3, random_state=42)
```

```python
from sklearn import tree

text_representation = export_text(model, feature_names= diabetes.feature_names)
print(text_representation)
```

Output:
```
|--- bmi <= 0.01
|   |--- s5 <= 0.01
|   |   |--- s5 <= -0.04
|   |   |   |--- value: [80.88]
|   |   |--- s5 >  -0.04
|   |   |   |--- value: [109.92]
|   |--- s5 >  0.01
|   |   |--- s4 <= 0.09
|   |   |   |--- value: [159.57]
|   |   |--- s4 >  0.09
|   |   |   |--- value: [256.33]
|--- bmi >  0.01
|   |--- bmi <= 0.07
|   |   |--- s6 <= 0.03
|   |   |   |--- value: [175.80]
|   |   |--- s6 >  0.03
|   |   |   |--- value: [230.52]
|   |--- bmi >  0.07
|   |   |--- s2 <= 0.02
|   |   |   |--- value: [291.22]
|   |   |--- s2 >  0.02
|   |   |   |--- value: [225.75]
```

```python
from sklearn.tree import _tree

def tree_to_code(tree, feature_names):

    tree_ = tree.tree_ # we are accessing the internal tree attribute 
    
    feature_name = [] # create a list to save the names of the features for each decision
    for i in tree_.feature: # instead of numbers we save the feature names that the nodes are split 
        if i != _tree.TREE_UNDEFINED: # if the node is not a leaf we append the value of the feature name that is split upon 
            feature_name.append(feature_names[i])
        else:
            feature_name.append("undefined!") # if it is a leaf we save it as undefined
    
    print("def predict({}):".format(", ".join(feature_names))) # based on which features we split 

    def recurse(node, depth):
        indent = "    " * depth
        if tree_.feature[node] != _tree.TREE_UNDEFINED: # if the node is not a leaf 
            name = feature_name[node] # save the name of the feature 
            threshold = tree_.threshold[node] # the threshold of the split 
            print("{}if {} <= {}:".format(indent, name, np.round(threshold,2))) 
            recurse(tree_.children_left[node], depth + 1) # go to left child 
            print("{}else:  # if {} > {}".format(indent, name, np.round(threshold,2)))
            recurse(tree_.children_right[node], depth + 1) # go to right child 
        else: # leaf node 
            print("{}return {}".format(indent, tree_.value[node]))

    recurse(0, 1)
```

```python
tree_to_code(model,  diabetes.feature_names)
```

Output:
```
def predict(age, sex, bmi, bp, s1, s2, s3, s4, s5, s6):
    if bmi <= 0.01:
        if s5 <= 0.01:
            if s5 <= -0.04:
                return [[80.87755102]]
            else:  # if s5 > -0.04
                return [[109.9223301]]
        else:  # if s5 > 0.01
            if s4 <= 0.09:
                return [[159.57407407]]
            else:  # if s4 > 0.09
                return [[256.33333333]]
    else:  # if bmi > 0.01
        if bmi <= 0.07:
            if s6 <= 0.03:
                return [[175.8]]
            else:  # if s6 > 0.03
                return [[230.51515152]]
        else:  # if bmi > 0.07
            if s2 <= 0.02:
                return [[291.22222222]]
            else:  # if s2 > 0.02
                return [[225.75]]
```

```python
def get_rules(tree, feature_names, class_names):
    tree_ = tree.tree_
    feature_name = [
        feature_names[i] if i != _tree.TREE_UNDEFINED else "undefined!"
        for i in tree_.feature
    ]

    paths = []
    path = []
    
    def recurse(node, path, paths):
        
        if tree_.feature[node] != _tree.TREE_UNDEFINED:
            name = feature_name[node]
            threshold = tree_.threshold[node]
            p1, p2 = list(path), list(path)
            p1 += [f"{name} <= {np.round(threshold, 2)}"]
            recurse(tree_.children_left[node], p1, paths)
            p2 += [f"{name} > {np.round(threshold, 2)}"]
            recurse(tree_.children_right[node], p2, paths)
        else:
            path += [(tree_.value[node], tree_.n_node_samples[node])]
            paths += [path]
            
    recurse(0, path, paths)

    # sort by samples count
    samples_count = [p[-1][1] for p in paths]
    ii = list(np.argsort(samples_count))
    paths = [paths[i] for i in reversed(ii)]
    
    rules = []
    for path in paths:
        rule = "αν "
        
        for p in path[:-1]:
            if rule != "αν ":
                rule += " και "
            rule += str(p)
        rule += " τότε "
        if class_names is None:
            rule += "response: "+str(np.round(path[-1][0][0][0],3))
        else:
            classes = path[-1][0][0]
            l = np.argmax(classes)
            rule += f"{class_names[l]}" #  (proba: {np.round(100.0*classes[l]/np.sum(classes),2)}%)
        #rule += f" | based on {path[-1][1]:,} samples"
        rules += [rule]
        
    return rules
```

```python


rules = get_rules(model,  diabetes.feature_names , None)
for r in rules:
    print(r)
```

Output:
```
αν bmi <= 0.01 και s5 <= 0.01 και s5 > -0.04 τότε response: 109.922
αν bmi > 0.01 και bmi <= 0.07 και s6 <= 0.03 τότε response: 175.8
αν bmi <= 0.01 και s5 > 0.01 και s4 <= 0.09 τότε response: 159.574
αν bmi <= 0.01 και s5 <= 0.01 και s5 <= -0.04 τότε response: 80.878
αν bmi > 0.01 και bmi <= 0.07 και s6 > 0.03 τότε response: 230.515
αν bmi > 0.01 και bmi > 0.07 και s2 <= 0.02 τότε response: 291.222
αν bmi > 0.01 και bmi > 0.07 και s2 > 0.02 τότε response: 225.75
αν bmi <= 0.01 και s5 > 0.01 και s4 > 0.09 τότε response: 256.333
```

```python
import numpy as np

def get_path(tree, sample , feature_names):
    tree_ = tree.tree_
    feature_name = [
        feature_names[i] if i != _tree.TREE_UNDEFINED else "undefined!"
        for i in tree_.feature
    ]

    path = []
    
    def recurse(node):
        nonlocal path
        
        if tree_.feature[node] != _tree.TREE_UNDEFINED:
            name = feature_name[node]
            threshold = tree_.threshold[node]
            
            # Check if the feature value of the sample satisfies the split condition
            if sample[feature_names.index(name)] <= threshold:
                path.append(f"{name} <= {np.round(threshold, 2)}")
                recurse(tree_.children_left[node])
            else:
                path.append(f"{name} > {np.round(threshold, 2)}")
                recurse(tree_.children_right[node])
        else:
            # Reached a leaf node, add the node to the path
            path.append((tree_.value[node], tree_.n_node_samples[node]))
            
    recurse(0)

    # Convert path to a readable rule
    rule = "αν "
    for p in path[:-1]:
        if rule != "αν ":
            rule += " και "
        rule += str(p)
    rule += " τότε "
    rule += str(np.round(path[-1][0][0][0], 2))
   
    return rule

# Usage example:
# Assuming X_test[0] is the first sample in your test data
sample_path = get_path(model, X_test[6], diabetes.feature_names)
print(sample_path)

```

Output:
```
αν bmi > 0.01 και bmi > 0.07 και s2 > 0.02 τότε 225.75
```

<!-- grade_id: cell-1697542afe8f99a2, points: 0, grade: false, solution: true -->
```python
def get_path(tree, instance, feature_names):
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
Cell In[101], line 3
    raise NotImplementedError()
    ^
IndentationError: expected an indented block after function definition on line 1
```

<!-- grade_id: cell-8aabb13a60a743dd, points: 4, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού του μονοπατιού"""
assert get_path(model, X_test[0], diabetes.feature_names) == 'αν bmi <= 0.01 και s5 > 0.01 και s4 <= 0.09 τότε 159.57'
assert get_path(model, X_test[1], diabetes.feature_names) == 'αν bmi > 0.01 και bmi <= 0.07 και s6 <= 0.03 τότε 175.8'
assert get_path(model, X_test[6], diabetes.feature_names) == 'αν bmi > 0.07 και s2 > 0.02 τότε 225.75'
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[103], line 4
      2 assert get_path(model, X_test[0], diabetes.feature_names) == 'αν bmi <= 0.01 και s5 > 0.01 και s4 <= 0.09 τότε 159.57'
      3 assert get_path(model, X_test[1], diabetes.feature_names) == 'αν bmi > 0.01 και bmi <= 0.07 και s6 <= 0.03 τότε 175.8'
----> 4 assert get_path(model, X_test[6], diabetes.feature_names) == 'αν bmi > 0.07 και s2 > 0.02 τότε 225.75'

AssertionError:
```

**5** Διαχωρίστε τα δεδομένα diabetes σε υποσύνολα εκπαίδευσης 60% (X_train, y_train), επικύρωσης 15% (X_val, y_val) και ελέγχου 25% (X_test, y_test) με σπόρο τυχαιότητας 42. (1 μονάδα)

<!-- grade_id: cell-c2448c38ed5250e1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
X_combine , X_test , y_combine , y_test = train_test_split( diabetes.data ,  diabetes.target , test_size= 0.25 , random_state= 42)

X_train , X_val , y_train , y_val = train_test_split( X_combine ,  y_combine , test_size= 0.20 , random_state= 42)

#print tests
# print(X_train[0][8])
# print(X_val[0][8])
# print(X_test[0][8])
#raise NotImplementedError()
```

<!-- grade_id: cell-8da00aed9ea7c071, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == -0.08024
assert round(X_val[0][8], 5) == 0.09924
assert round(X_test[0][8], 5) == 0.03243
```

**6** Υλοποιήστε συνάρτηση, η οποία δοθέντος ενός δενδρικού μοντέλου παλινδρόμησης, διερευνά όλες τις ακέραιες τιμές της μεταβλητής min_samples_leaf από ένα κάτω όριο μέχρι ένα άνω όριο και επιστρέφει την βέλτιστη με βάση το mean squared error στο παραπάνω σύνολο επικύρωσης έπειτα από εκπαίδευση στο παραπάνω σύνολο εκπαίδευσης. Επιπλέον επιστρέφει και το mean squared error (στρογγυλοποιημένο σε ακέραιο) που πετυχαίνει στο σύνολο ελέγχου το μοντέλο με την βέλτιστη επίδοση στο σύνολο επικύρωσης. (2 μονάδες)

<!-- grade_id: cell-45e2ac4f9c5115f9, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import mean_squared_error

def best_min_samples_leaf(model, min, max):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0d6be21e54be9665, points: 2, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=3, random_state=42)
assert best_min_samples_leaf(model, 1, 8) == (7, 3358)
assert best_min_samples_leaf(model, 5, 15) == (11, 3271)
```

**7** Υλοποιήστε συνάρτηση η οποία δοθέντος ενός μοντέλου παλινδρόμησης, ενός συνόλου δεδομένων και μίας μετρικής, το αξιολογεί με τη μέθοδο της σταυρωτής επικύρωσης "άφησε ένα εκτός". (1 μονάδα)

<!-- grade_id: cell-880c64392f85bccb, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import LeaveOneOut
import numpy as np

def leave_one_out(model, X, y, metric):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4afd7b6de1cd1f7f, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_leaf_nodes=6, min_samples_leaf=2, random_state=42)
assert round(leave_one_out(model, diabetes.data, diabetes.target, mean_squared_error), 0) == 3354
```